```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

# Homework 6: Probability, Simulation, Estimation, and Assessing Models

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 
* [Randomness](https://www.inferentialthinking.com/chapters/09/Randomness.html)
* [Sampling and Empirical Distributions](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Testing Hypotheses](https://www.inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

## 1. Sic Bo

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')
```

Sic Bo is a popular dice game often found in casinos, where players wager on the outcome of the roll of three six-sided dice. If a player's combined outcome of each die amounts to any number between 4 and 10, the player is said to have a **SMALL**. For a combined outcome between 11 and 17, the player is said to have a **BIG**. 

In our simulation, inspired by Sic Bo, players can place bets on either a **SMALL** or **BIG** and win accordingly. The table `rolls` shows all the possible combinations of 3 dice rolls and their respective outcomes.

Run the cell below to load the `rolls` table.

```python
rolls = Table.read_table('rolls.csv').drop(0)
rolls
```

### Betting on SMALL ###
If you bet on **SMALL**, you are betting that the outcome of rolling the dice will be a number between 4 and 10. This bet *pays 1 to 1*. That means if you place a one-dollar bet on a **SMALL**, then:

- If the outcome of the dice roll is between 4 and 10, your gain is 1 dollar. 
- If the outcome of the dice roll is not between 4 and 10, you lose your dollar. In other words, your gain is -1 dollars.

Let's see if you can make money by betting on a **SMALL** in Sic Bo.

**Question 1.** Define a function `dollar_bet_on_small` that takes the result of the outcome of a roll and returns your gain in dollars if the outcome is a **SMALL** and you had placed a one-dollar bet on **SMALL**. Remember that the gain can be negative. Make sure your function returns an integer. **(4 points)**

*Note:* You can assume that the arguments passed to the function will only be numbers from the `Outcome` column. Your function doesn't have to check  for this condition.


```python
def dollar_bet_on_small(...):
    ...
```

```python
grader.check("q1_1")
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_small(4))
print(dollar_bet_on_small(9))
print(dollar_bet_on_small(18))
```

**Question 2.** Add a column labeled `Winnings: Small` to the table `rolls`. For each outcome, the column should contain your gain in dollars if that outcome won and you had bet one dollar on **SMALL**. Your code should use the function `dollar_bet_on_small`. **(4 points)**


```python
small_winnings = ...
rolls = ...
rolls
```

```python
grader.check("q1_2")
```

### Simulating 10 Bets on SMALL
For each dice roll, the winning outcome is equally likely to be any of the 216 outcomes regardless of the results of all other rolls. Let's see what would happen if we decided to bet one dollar on **SMALL** each round.

**Question 3.** Create a table named `ten_bets` by sampling the rolls table to simulate 10 rolls of three dice. Ensure that your table has the same four column labels as in `rolls`. Once you've created that table, set `sum_bets` to your net gain in all 10 bets, assuming that you bet one dollar on **SMALL** each time. **(4 points)**

*Hint:* It may be helpful to print out `ten_bets` after you create it!


```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Small'))
sum_bets
```

```python
grader.check("q1_3")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on **SMALL**. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

**Question 4.** Let's see what would happen if you made more bets. Define a function `net_gain_small` that takes the number of bets and returns the net gain in that number of one-dollar bets on **SMALL**. **(4 points)**

*Hint:* You should use the `rolls` table within your function.


```python
def net_gain_small(...):
    ...
```

```python
grader.check("q1_4")
```

Run the cell below a few times to make sure that the results are similar to those you observed in the previous exercise.

```python
net_gain_small(10)
```

**Question 5.** Complete the cell below to simulate the net gain in 200 one-dollar bets on **SMALL**, repeating the process 10,000 times. After the cell is run, `all_gains_small` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on **SMALL**. **(4 points)**

*Hint:* Think about which computational tool might be helpful for simulating a process multiple times. Lab 5 might be a good resource to look at!

*Note:* This cell might take a few seconds to run.


```python
num_bets = ...
repetitions = ...

all_gains_small = ...
...

len(all_gains_small) # Do not change this line! Check that all_gains_small is length 10000.
```

```python
grader.check("q1_5")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Small', all_gains_small)
gains.hist(bins=np.arange(-80, 70, 10), edgecolor='black') 

```

**Question 6:** Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on a small, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question. **(4 points)**


```python
loss_more_than_50 = ...
```

```python
grader.check("q1_6")
```

### Betting on a Single ###
Let's try a different type of bet. In our simulated game, a bet on a *single* is a wager that a specific number will appear on a specific one of the three dice. First, specify one of the three dice. Then choose a single number from 1 to 6, and if your number appears on your chosen die, the payout is 5 to 1. For example, if you place a one-dollar bet on Die 1 landing on 6, then:

- If Die 1 lands on a 6, your gain is 5 dollars.
- If Die 1 does not land on a 6, you lose your dollar, so your gain is -1 dollars.

**Question 7.** Define a function called `dollar_bet_on_six` that takes a specific die number from 1 to 6 as input, and calculates your gain in dollars (either 5 dollars or -1 dollars). The function should determine the outcome of betting a dollar on 6.  **(4 points)**

*Note*: We will focus on the outcome of rolling only a single die for the purpose of this exercise.

```python
def dollar_bet_on_six(...):
    ...
```

```python
grader.check("q1_7")
```

Run the cell below to check that your function is doing what it should.

```python
print(dollar_bet_on_six(3))
print(dollar_bet_on_six(4))
print(dollar_bet_on_six(6))
print(dollar_bet_on_six(5))
```

**Question 8.** Add a column `Winnings: Single` to the `rolls` table. Pick a single die (either Die 1, Die 2 or Die 3). For each bet (i.e., roll of the 3 dice), the `Winnings: Single` column should contain your gain in dollars if the chosen Die lands on a 6 and you had bet one dollar. **(4 points)**

*Note*: Each die is equally likely to land on the number 6. The Die you choose does not impact the outcome.


```python
# Any "Die" column is correct
single_winnings = ...
rolls = ...
rolls.show(5) # Do not change this line.
```

```python
grader.check("q1_8")
```

**Question 9.** Simulate the net gain in 200 one-dollar bets on the single 6, repeating the process 10,000 times and saving your gains in the array `all_gains_single`. **(4 points)**

*Hint:* Your code in Questions 4 and 5 may be helpful here!


```python
all_gains_single = ...

...

# Do not change the two lines below
gains = gains.with_columns('Net Gain on Single', all_gains_single)
gains.hist(bins = np.arange(-200, 150, 20))
```

```python
grader.check("q1_9")
```

**Question 10.** Look carefully at the visualization above, and assign `histogram_statements` to an array of the numbers of each statement below that can be correctly inferred from the overlaid histogram. **(4 points)**

1. If you bet one dollar 200 times on a small, your chance of loosing money is more than 50%.
2. If you bet one dollar 200 times in our simulaton, your chance of loosing more than 50 dollars is greater if you bet on a small each time than if you bet on a single each time.
3. If you bet one dollar 200 times in our simulaton, your chance of making more than 50 dollars is greater if you bet on a single each time than if you bet on small each time.

*Hint:* We've already seen one of these statements in a prior question.


```python
histogram_statements = ...
```

```python
grader.check("q1_10")
```

## 2. Chances

Before you do this exercise, make sure you understand the logic behind all the examples in [Section 9.5](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen.
- Seeing if it is easier to find the chance that the event does not happen.

### Finding Chances

When rolling 3 dice, there are 6 possible outcomes for each die you roll. These outcomes are equally likely regardless of the results of other rolls. Among these outcomes, various combinations lead to different results. In each section below, provide an expression that calculates the probability of the specified event.

**Question 1.** Rolling 3 fives on your first try. **(4 points)**

*Hint:* We are rolling 3 dice at time.


```python
first_roll_all_five = ...
```

```python
grader.check("q2_1")
```

**Question 2.** Not rolling 3 fives on your first try. **(4 points)**


```python
no_five = ...
```

```python
grader.check("q2_2")
```

**Question 3.** At least one of the 3 dice rolled is a four. **(4 points)**


```python
at_least_one_four = ...
at_least_one_four
```

```python
grader.check("q2_3")
```

**Question 4.** You don't roll a double or a triple. In other words, all three dice have unique numbers. **(4 points)**

*Hint:* Consider how many possible outcomes there and how many of them are doubles or triples. Note for example, that rolling a 1, a 2, and a 1 consistitutes a double since 2 dice are the same. 


```python
lone_winners = ...
```

```python
grader.check("q2_4")
```

### Comparing Chances
In each of Questions 5-7, two events A and B are described. Choose from one of the following three options and set each answer variable to a single integer:

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time, and by thinking about the [law of averages](https://inferentialthinking.com/chapters/10/1/Empirical_Distributions.html#the-law-of-averages).

**Question 5.** A gardener selects three times at random from a basket containing three different types of flowers: roses, tulips, and daisies. **(4 points)**

- Event A: all three different flowers are picked (assuming the gardener picks with replacement)
- Event B: all three different flowers are picked (assuming the gardener picks without replacement)


```python
flowers_option = ...
```

```python
grader.check("q2_5")
```

**Question 6.** In a raffle, two tickets are drawn at random without replacement from a bag containing tickets numbered 1 through 500. **(4 points)**

- Event A: The numbers 500 and 1 are picked on the first and second draws, respectively
- Event B: The numbers 10 and 15 are picked on the first and second draws, respectively


```python
raffle_option = ...
```

```python
grader.check("q2_6")
```

**Question 7.** A fair coin is tossed repeatedly. **(4 points)**

- Event A: There are 40 or less heads in 100 tosses
- Event B: There are 400 or less heads in 1000 tosses

*Hint*: Think about the law of averages!


```python
coin_option = ...
```

```python
grader.check("q2_7")
```

## 3. Three Ways Python Draws Random Samples

You have learned three ways to draw random samples using Python:

- `tbl.sample` draws a random sample of rows from the table `tbl`. The output is a table consisting of the sampled rows. 

- `np.random.choice` draws a random sample from a population whose elements are in an array. The output is an array consisting of the sampled elements.

- `sample_proportions` draws from a categorical distribution whose proportions are in an array. The output is an array consisting of the sampled proportions in all the categories. 

```python
# Look through this code and run this cell for questions 1 and 2
top = Table.read_table('top_movies_2017.csv').select(0, 1)
top.show(3)
```

```python
# Look through this code and run this cell for questions 1 and 2
studios_with_counts = top.group('Studio').sort('count', descending=True)
studios_with_counts.show(3)
```

```python
# Look through this code and run this cell for questions 1 and 2
studios_of_all_movies = top.column('Studio')
distinct_studios = studios_with_counts.column('Studio')

print("studios_of_all_movies:", studios_of_all_movies[:10], "...")
print("\n distinct_studios:", distinct_studios)
```

```python
# Look through this code and run this cell for questions 1 and 2
studio_counts_only = studios_with_counts.column('count')
studio_proportions_only = studio_counts_only / sum(studio_counts_only)

print("studio_counts_only:", studio_counts_only)
print("\n studio_proportions_only:", studio_proportions_only)
```

In Questions 1 and 2, we will present a scenario. Determine which three of the following six options are true in regards to what the question is asking, and list them in the following answer cell. If your answer includes any of (i)-(iii), state what you would fill in the blank to make it true: `top`, `studios_with_counts`, `studios_of_all_movies`, `distinct_studios`, `studio_counts_only` or `studio_proportions_only`.

(i) This can be done using `sample` and the table _________.

(ii) This can be done using `np.random.choice` and the array ________.

(iii) This can be done using `sample_proportions` and the array _______.

(iv) This cannot be done using `sample` and the data given.

(v) This cannot be done using `np.random.choice` and the data given.

(vi) This cannot be done using `sample_proportions` and the data given.

<!-- BEGIN QUESTION -->

**Question 1.** Simulate a sample of 10 movies drawn at random with replacement from the 200 movies. Outputs True if the same studio released both the first and last sampled movie. **(4 points)**

*Example Answer:* (i) studios_of_all_movies, (iii) top, (v)

***Note***: Do not explain your answer for any of the options you've chosen; please follow the structure of the example answer provided.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.** Simulate a sample of 10 movies drawn at random with replacement from the 200 movies. Outputs True if Universal appears more often than Disney among studios that released the sampled movies, and False otherwise. **(4 points)**

*Example Answer:* (i) studios_of_all_movies, (iii) top, (v)

***Note***: Do not explain your answer for any of the options you've chosen; please follow the structure of the example answer provided.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Assessing Jade's Models

Before you begin, [Section 10.4](https://inferentialthinking.com/chapters/10/4/Random_Sampling_in_Python.html) of the textbook is a useful reference for this part.

### Games with Jade
Our friend Jade comes over and asks us to play a game with her. The game works like this: 

> We will draw randomly with replacement from a simplified 13 card deck with 4 face cards (A, J, Q, K), and 9 numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10). If we draw cards with replacement 13 times, and if the number of numbered cards is greater than or equal to 9, we win.
> 
> Otherwise, Jade wins.

We play the game once and we lose, observing 5 total numbered cards. We are angry and accuse Jade of cheating! Jade is adamant, however, that the deck is fair.

Jade's model claims that there is an equal chance of getting any of the cards (A, 2, 3, 4, 5, 6, 7, 8, 9, 10, J, Q, K), but we do not believe her. We believe that the deck is clearly rigged, with numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10) being less likely than the face cards (A, J, Q, K).

**Question 1.** Assign `deck_model_probabilities` to a two-item array containing the chance of drawing a face card as the first element, and the chance of drawing a numbered, non-face card as the second element under Jade's model. Since we're working with probabilities, make sure your values are between 0 and 1. **(4 Points)**


```python
deck_model_probabilities = ...
deck_model_probabilities
```

```python
grader.check("q4_1")
```

**Question 2.** We believe Jade's model is incorrect. In particular, we believe there to be a *smaller* chance of getting a numbered, non-face card. Which of the following statistics can we use during our simulation to test between the model and our alternative? Assign `statistic_choice` to the correct answer. **(5 Points)**

1. The distance (absolute value) between the actual number of non-face cards in 13 draws and 9, the expected number of non-face cards in 13 draws
2. The number of non-face cards we get in 13 draws
3. The expected number of non-face cards in 13 draws, which is 9


```python
statistic_choice = ...
statistic_choice
```

```python
grader.check("q4_2")
```

**Question 3.** Define the function `deck_simulation_and_statistic`, which, given a sample size and an array of model proportions (like the one you created in Question 1), returns the number of numbered cards in one simulation of drawing cards under the model specified in `model_proportions`. **(5 Points)**

*Hint:* Think about how you can use the function `sample_proportions`. 


```python
def deck_simulation_and_statistic(sample_size, model_proportions):
    ...

deck_simulation_and_statistic(13, deck_model_probabilities)
```

```python
grader.check("q4_3")
```

**Question 4.** Use your function from above to simulate the drawing of 13 cards 2500 times under the proportions that you specified in Question 1. Keep track of all of your statistics in `deck_statistics`. **(5 Points)**


```python
repetitions = 2500 
...

deck_statistics
```

```python
grader.check("q4_4")
```

Let’s take a look at the distribution of simulated statistics.

```python
# Draw a distribution of statistics 
Table().with_column('Deck Statistics', deck_statistics).hist()
```

<!-- BEGIN QUESTION -->

**Question 5.** Given your observed value, do you believe that Jade's model is reasonable, or is our alternative (that our deck is rigged) more likely? Explain your answer using the histogram produced above. **(5 Points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1_1 (points: [2])
```
>>> dollar_bet_on_small(6) == 1
True
```


Tests for q1_10 (points: [0, 0])
```
>>> import numpy as np
>>> type(histogram_statements) == np.ndarray
True
```
```
>>> 1 <= histogram_statements.item(0) <= 3
True
```


Tests for q1_2 (points: [0])
```
>>> rolls.labels[4] == 'Winnings: Small'
True
```


Tests for q1_3 (points: [0])
```
>>> set(['Outcome', 'Die 1', 'Die 2', 'Die 3', 'Winnings: Small']) == set(ten_bets.labels)
True
```


Tests for q1_4 (points: [0])
```
>>> net_gain_small(10000) != net_gain_small(10000)
True
```


Tests for q1_5 (points: [0])
```
>>> len(all_gains_small) == 10000
True
```


Tests for q1_6 (points: [0])
```
>>> type(loss_more_than_50) == bool
True
```


Tests for q1_7 (points: [2])
```
>>> dollar_bet_on_six(6) == 5
True
```


Tests for q1_8 (points: [0, 4])
```
>>> rolls.labels[5] == 'Winnings: Single'
True
```
```
>>> max(rolls.column('Winnings: Single')) == 5
True
```


Tests for q1_9 (points: [0])
```
>>> len(all_gains_single) == 10000
True
```


Tests for q2_1 (points: [0])
```
>>> 0 < first_roll_all_five < 1
True
```


Tests for q2_2 (points: [0])
```
>>> 0 < no_five < 1
True
```


Tests for q2_3 (points: [0])
```
>>> 0 < at_least_one_four < 1
True
```


Tests for q2_4 (points: [0])
```
>>> 0 < lone_winners < 1
True
```


Tests for q2_5 (points: [0, 0])
```
>>> type(flowers_option) == int
True
```
```
>>> 1 <= flowers_option <= 3
True
```


Tests for q2_6 (points: [0, 0])
```
>>> type(raffle_option) == int
True
```
```
>>> 1 <= raffle_option <= 3
True
```


Tests for q2_7 (points: [0, 0])
```
>>> type(coin_option) == int
True
```
```
>>> 1 <= coin_option <= 3
True
```


Tests for q4_1 (points: [0, 0])
```
>>> len(deck_model_probabilities) == 2
True
```
```
>>> sum(deck_model_probabilities) == 1
True
```


Tests for q4_2 (points: [0, 0])
```
>>> type(statistic_choice) == int
True
```
```
>>> 1 <= statistic_choice <= 3
True
```


Tests for q4_3 (points: [5])
```
>>> num_face = deck_simulation_and_statistic(13, deck_model_probabilities)
>>> 0 <= num_face <= 13
True
```


Tests for q4_4 (points: [2, 3])
```
>>> len(deck_statistics) == repetitions
True
```
```
>>> all([0 <= k <= 13 for k in deck_statistics])
True
```
