```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# Lab 4: Functions and Visualizations

Welcome to Lab 4! This week, we'll learn about functions, table methods such as `apply`, and how to generate visualizations! 

Recommended Reading:

* [Applying a Function to a Column](https://inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Visualizations](https://inferentialthinking.com/chapters/07/Visualization.html)
* [Python Reference](http://data8.org/sp22/python-reference.html)

**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, and then execute the final cell. Then submit the downloaded zip file, that includes your notebook,  according to your instructor's directions.

First, set up the notebook by running the cell below.

```python
import numpy as np
from datascience import *

# These lines set up graphing capabilities.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')

from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets
```

## 1. Defining functions

Let's write a very simple function that converts a proportion to a percentage by multiplying it by 100.  For example, the value of `convert_to_percentage(.5)` should be the number 50 (no percent sign).

A function definition has a few parts.

##### `def`
It always starts with `def` (short for **def**ine):

    def

##### Name
Next comes the name of the function.  Like other names we've defined, it can't start with a number or contain spaces. Let's call our function `convert_to_percentage`:
    
    def convert_to_percentage

##### Signature
Next comes something called the *signature* of the function.  This tells Python how many arguments your function should have, and what names you'll use to refer to those arguments in the function's code.  A function can have any number of arguments (including 0!). 

`convert_to_percentage` should take one argument, and we'll call that argument `proportion` since it should be a proportion.

    def convert_to_percentage(proportion)
    
If we want our function to take more than one argument, we add a comma between each argument name. Note that if we had zero arguments, we'd still place the parentheses () after that name. 

We put a **colon** after the signature to tell Python that the next indented lines are the body of the function. If you're getting a syntax error after defining a function, check to make sure you remembered the colon!

    def convert_to_percentage(proportion):

##### Documentation
Functions can do complicated things, so you should write an explanation of what your function does.  For small functions, this is less important, but it's a good habit to learn from the start.  Conventionally, Python functions are documented by writing an **indented** triple-quoted string:

    def convert_to_percentage(proportion):
        """Converts a proportion to a percentage."""

##### Body
Now we start writing code that runs when the function is called.  This is called the *body* of the function and every line **must be indented with a tab**.  Any lines that are *not* indented and left-aligned with the def statement is considered outside the function. 

Some notes about the body of the function:
- We can write code that we would write anywhere else.  
- We use the arguments defined in the function signature. We can do this because we assume that when we call the function, values are already assigned to those arguments.
- We generally avoid referencing variables defined *outside* the function. If you would like to reference variables outside of the function, pass them through as arguments!


Now, let's give a name to the number we multiply a proportion by to get a percentage:

    def convert_to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100

##### `return`
The special instruction `return` is part of the function's body and tells Python to make the value of the function call equal to whatever comes right after `return`.  We want the value of `convert_to_percentage(.5)` to be the proportion .5 times the factor 100, so we write:

    def convert_to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100
        return proportion * factor
        
`return` only makes sense in the context of a function, and **can never be used outside of a function**. `return` is always the last line of the function because Python stops executing the body of a function once it hits a `return` statement. If a function does not have a return statement, it will not return anything; if you expect a value back from the function, make sure to include a return statement. 

*Note:*  `return` inside a function tells Python what value the function evaluates to. However, there are other functions, like `print`, that have no `return` value. For example, `print` simply prints a certain value out to the console. 

In short, `return` is used when you want to tell the *computer* what the value of some variable is, while `print` is used to tell you, a *human*, its value.

**Question 1.1.** Define `convert_to_percentage` in the cell below.  Call your function to convert the proportion .3 to a percentage.  Name that percentage `thirty_percent`.


```python
def ...
    ''' ... '''
    ... = ...
    return ...

thirty_percent = ...
thirty_percent
```

```python
grader.check("q11")
```

Here's something important about functions: the names assigned *within* a function body are only accessible within the function body. Once the function has returned, those names are gone.  So even if you created a variable called `factor` and defined `factor = 100` inside of the body of the `convert_to_percentage` function and then called `convert_to_percentage`, `factor` would not have a value assigned to it outside of the body of `convert_to_percentage`:

**Note**: Below, you should see an error message starting with "Uh-o", along with the official error from Python under it. The first message that you see is specifically from Data 8 Staff to provide extra debugging help. It will not appear if you run into Python errors outside of Data 8.

```python
# You should get an error when you run this.  (If you don't, 
# you might have defined factor somewhere above.)
factor
```

Like you’ve done with built-in functions in previous labs (max, abs, etc.), you can pass in named values as arguments to your function.

**Question 1.2.** Use `convert_to_percentage` again to convert the proportion named `a_proportion` (defined below) to a percentage called `a_percentage`.

*Note:* You don't need to define `convert_to_percentage` again!  Like other named values, functions stick around after you define them.


```python
a_proportion = 3**(0.5) / 3
a_percentage = ...
a_percentage
```

```python
np.isclose(a_percentage, 57.735026918962575)
```

As we've seen with built-in functions, functions can also take strings (or arrays, or tables) as arguments, and they can return those things, too.

In the following cell, we will define a function called `change_string`.  It takes in a single string as its argument. It returns a copy of that string, but with all the characters that are spaces and the letter "a" removed.  

To remove all letter "a"s from a string, we used `a_string.replace("a", "")`.  The `.replace` method for strings returns a new string, so we can call `replace` multiple times, one after the other. 

```python
def change_string(a_string):
    """Removes all spaces and the letter 'a' from a string."""
    return a_string.replace(" ", "").replace("a","")

# An example call to the function.  (It's often helpful to run
# an example call from time to time while we're writing a function,
# to see how it currently works.)
change_string("Can you read this without spaces and the letter a?")
```

##### Calls on calls on calls
Just as you write a series of lines to build up a complex computation, it's useful to define a series of small functions that build on each other.  Since you can write any code inside a function's body, you can call other functions you've written.

If a function is a like a recipe, defining a function in terms of other functions is like having a recipe for cake telling you to follow another recipe to make the frosting, and another to make the jam filling.  This makes the cake recipe shorter and clearer, and it avoids having a bunch of duplicated frosting recipes.  It's a foundation of productive programming.

For example, suppose you want to count the number of characters *that aren't spaces or the letter "a"* in a piece of text.  One way to do that is this to remove all the spaces and the letter "a"s and count the size of the remaining string.

**Question 1.3.** Write a function called `num_string_characters`.  It should take a string as its argument and return a number.  That number should be the number of characters in the argument string that aren't spaces or the letter "a". You should use the `change_string` function we provided above inside of the `num_string_characters` function.

*Hint:* The function `len` takes a string as its argument and returns the number of characters in it.


```python
def num_string_characters(a_string):
    """The number of characters in a string, minus spaces and the letter a."""
    ...

# Try calling your function yourself to make sure the output is what
# you expect. 
num_string_characters("Go Laney!")
```

```python
grader.check("q13")
```

Functions can also encapsulate code that *displays output* instead of computing a value. For example, if you call `print` inside a function, and then call that function, something will get printed.

The `movies_by_year` dataset in the textbook has information about movie sales in recent years.  Suppose you'd like to display the name of the movie with the 5th-highest total gross movie sales, printed within a sentence.  You might do this:

```python
movies_by_year = Table.read_table("movies_by_year.csv")
rank = 5
fifth_from_top_movie_name = movies_by_year.sort("Total Gross", descending=True).column("#1 Movie").item(rank-1)
print("The movie at rank", rank, "for total gross movie sales was:", fifth_from_top_movie_name)
```

After writing this, you realize you also wanted to print out the names of the 2nd and 3rd-highest grossing movies.  Instead of copying your code, you decide to put it in a function.  Since the rank varies, you make that an argument to your function.

**Question 1.4.** Write a function called `print_kth_top_movie_name`.  It should take a single argument, the rank of the year (like 2, 3, or 5 in the above examples) and should use the table `movies_by_year`. It should print out a message like the one above.  

*Note:* Your function shouldn't have a `return` statement.


```python
def print_kth_top_movie_name(k):
    ...
    ...
    print(...)

# Example calls to your function:
print_kth_top_movie_name(2)
print_kth_top_movie_name(3)
```

```python
grader.check("q14")
```

```python
# interact also allows you to pass in an array for a function argument. It will
# then present a dropdown menu of options.
_ = interact(print_kth_top_movie_name, k=np.arange(1, 10))
```

### `print` is not the same as `return`
The `print_kth_top_movie_name(k)` function prints the name of the movie with total gross movie sales for the rank that was provided! However, since we did not return any value in this function, we can not use it after we call it. Let's look at an example of another function that prints a value but does not return it.

```python
def print_number_five():
    print(5)
```

```python
print_number_five()
```

However, if we try to use the output of `print_number_five()`, we see that the value `5` is printed but we get a TypeError when we try to add the number 2 to it!

```python
print_number_five_output = print_number_five()
print_number_five_output + 5
```

It may seem that `print_number_five()` is returning a value, 5. In reality, it just displays the number 5 to you without giving you the actual value! If your function prints out a value **without returning it** and you try to use that value, you will run into errors, so be careful!

Explain to your neighbor or a staff member how you might add a line of code to the `print_number_five` function (after `print(5)`) so that the code `print_number_five_output + 5` would result in the value `10`, rather than an error.

## 2. Functions and CEO Incomes

In this question, we'll look at the 2015 compensation of CEOs at the 100 largest companies in California. The data was compiled from a [Los Angeles Times analysis](http://spreadsheets.latimes.com/california-ceo-compensation/), and ultimately came from [filings](https://www.sec.gov/answers/proxyhtf.htm) mandated by the SEC from all publicly-traded companies. Two companies have two CEOs, so there are 102 CEOs in the dataset.

We've copied the raw data from the LA Times page into a file called `raw_compensation.csv`. (The page notes that all dollar amounts are in **millions of dollars**.)

```python
raw_compensation = Table.read_table('raw_compensation.csv').sort("Cash Pay")
raw_compensation
```

Companies pay executives in a variety of ways: in cash, by granting stock or other equity in the company, or with ancillary benefits (like private jets). We want to compute the average of the CEOs' cash pay. Try running the cell below.

```python
np.average(raw_compensation.column("Cash Pay"))
```

You should see a `TypeError`. Let's examine why this error occurred by looking at the values in the `Cash Pay` column. To do so, we can use the `type` function. This function tells us the data type of the object that we pass into it. Run the following cell to see what happens when we pass in `23` to the `type` function. Does the result make sense?

```python
type(23)
```

**Question 2.1.** Use the `type` function and set `cash_pay_type` to the type of the first value in the "Cash Pay" column.


```python
cash_pay_type = ...
cash_pay_type
```

```python
grader.check("q21")
```

**Question 2.2.** You should have found that the values in the `Cash Pay` column are strings. It doesn't make sense to take the average of string values, so we need to convert them to numbers if we want to do this. Extract the tenth value in `Cash Pay`.  It's Marissa Mayer's pay in 2015, in *millions* of dollars.  Call it `marissa_mayer_pay_string`.


```python
marissa_mayer_pay_string = ...
marissa_mayer_pay_string
```

```python
grader.check("q22")
```

**Question 2.3.** Convert `marissa_mayer_pay_string` to a number of *dollars*. 

Some hints, as this question requires multiple steps:
- The string method `strip` will be useful for removing the dollar sign; it removes a specified character from the start or end of a string.  For example, the value of `"100%".strip("%")` is the string `"100"`.  
- You'll also need the function `float`, which converts a string that looks like a number to an actual number. Don't worry about the whitespace at the end of the string; the `float` function will ignore this.
- Finally, remember that the answer should be in dollars, not millions of dollars.


```python
marissa_mayer_pay = ...
marissa_mayer_pay
```

```python
grader.check("q23")
```

To compute the average cash pay, we need to do this for every CEO.  But that looks like it would involve copying this code 102 times.

This is where functions come in.  First, we'll define a new function, giving a name to the expression that converts "cash pay" strings to numeric values.  Later in this lab, we'll see the payoff: we can call that function on every pay string in the dataset at once.

The next section of this lab we will get some more practice on defining functions. For this part, just fill in the ellipses in the cell below.

**Question 2.4.** Copy the expression you used to compute `marissa_mayer_pay`, and use it as the return expression of the function below. But make sure you replace the specific `marissa_mayer_pay_string` with the generic `pay_string` name specified in the first line in the `def` statement.

*Hint*: When dealing with functions, you should generally not be referencing any variable outside of the function. Usually, you want to be working with the arguments that are passed into it, such as `pay_string` for this function. If you're using `marissa_mayer_pay_string` within your function, you're referencing an outside variable!


```python
def convert_pay_string_to_number(pay_string):
    """Converts a pay string like '$100' (in millions) to a number of dollars."""
    ...
```

```python
grader.check("q24")
```

Running that cell doesn't convert any particular pay string. Instead, it creates a function called `convert_pay_string_to_number` that can convert *any* string with the right format to a number representing millions of dollars.

We can call our function just like we call the built-in functions we've seen. It takes one argument, *a string*, and it returns a float.

```python
convert_pay_string_to_number('$42')
```

```python
convert_pay_string_to_number(marissa_mayer_pay_string)
```

```python
# We can also compute Safra Catz's pay in the same way:
convert_pay_string_to_number(raw_compensation.where("Name", are.containing("Safra")).column("Cash Pay").item(0))
```

So, what have we gained by defining the `convert_pay_string_to_number` function? 
Well, without it, we'd have to copy the code `10**6 * float(some_pay_string.strip("$"))` each time we wanted to convert a pay string.  Now we just call a function whose name says exactly what it's doing.

## 3. `apply`ing functions

Defining a function is a lot like giving a name to a value with `=`.  In fact, a function is a value just like the number 1 or the text "data"!

For example, we can make a new name for the built-in function `max` if we want:

```python
our_name_for_max = max
our_name_for_max(2, 6)
```

The old name for `max` is still around:

```python
max(2, 6)
```

Try just writing `max` or `our_name_for_max` (or the name of any other function) in a cell, and run that cell.  Python will print out a (very brief) description of the function.

```python
max
```

Now try writing `max?` or `our_name_for_max?` (or the name of any other function) in a cell, and run that cell.  A information box should show up at the bottom of your screen a longer description of the function

*Note: You can also press Shift+Tab after clicking on a name to see similar information!*

```python
our_name_for_max?
```

Let's look at what happens when we set `max`to a non-function value. You'll notice that a TypeError will occur when you try calling `max`. Things like integers and strings are not callable like this: `"hello"(2)`, which is calling a string and will error. Look out for any functions that might have been renamed when you encounter this type of error

```python
max = 6
max(2, 6)
```

```python
# This cell resets max to the built-in function. Just run this cell, don't change its contents
import builtins
max = builtins.max
```

Why is this useful?  Since functions are just values, it's possible to pass them as arguments to other functions.  Here's a simple but not-so-practical example: we can make an array of functions.

```python
make_array(max, np.average, are.equal_to)
```

**Question 3.1.** Make an array containing any 4 other functions you've seen.  Call it `four_functions`.


```python
four_functions = ...
four_functions
```

```python
grader.check("q31")
```

Working with functions as values can lead to some funny-looking code. For example, see if you can figure out why the following code works. Check your explanation with a peer or a staff member.

```python
make_array(max, np.average, are.equal_to).item(0)(4, -2, 7)
```

A more useful example of passing functions to other functions as arguments is the table method `apply`.

`apply` calls a function many times, once on *each* element in a column of a table.  It produces an *array* of the results.  Here we use `apply` to convert every CEO's total pay to a number, using the function you defined:

*Note:* You'll see a number like `5.325e+07`. This is Python's way of representing **scientific notation**. We interpret `5.325e+07` as `5.325 * 10**7`, or `53,250,000`.

```python
raw_compensation.apply(convert_pay_string_to_number, "Total Pay")
```

Here's an illustration of what that did:

<img src="apply.png"/>

Note that we didn’t write `raw_compensation.apply(convert_pay_string_to_number(), “Total Pay”)` or `raw_compensation.apply(convert_pay_string_to_number(“Total Pay”))`. We just passed the name of the function, with no parentheses, to `apply`, because all we want to do is let `apply` know the name of the function we’d like to use and the name of the column we’d like to use it on. `apply` will then call the function `convert_pay_string_to_number` on each value in the column for us!

**Question 3.2.** Using `apply`, make a table that's a copy of `raw_compensation` with one additional column called `Cash Pay ($)`.  That column should contain the result of applying `convert_pay_string_to_number` to the `Cash Pay` column (similar to the way we applied that function to the 'Total Pay' column above).  Call the new table `cash_compensation`.


```python
cash_compensation = raw_compensation.with_column(
    "Cash Pay ($)",
    ...
    )
cash_compensation
```

```python
grader.check("q32")
```

Now that we have all the pays as numbers, we can learn more about them through computation.

**Question 3.3.** Compute the average cash pay of the CEOs in the dataset.


```python
average_cash_pay = ...
average_cash_pay
```

```python
grader.check("q33")
```

**Question 3.4** As stated above, companies pay executives in a variety of ways: in cash, by granting stock or other equity in the company, or with ancillary benefits (like private jets).  Compute the proportion of each CEO's total pay that was equity.  (Your answer should be an array of numbers, one for each CEO in the dataset.) 


*Hint:* What function have you defined that can convert a string to a number?


```python
equity_proportion = ...
equity_proportion
```

```python
grader.check("q34")
```

**Why is `apply` useful?**

For operations like arithmetic, or the functions in the NumPy library, you don't need to use `apply`, because they automatically work on each element of an array.  But there are many things that don't.  The string manipulation we did in today's lab is one example.  Since you can write any code you want in a function, `apply` gives you total control over how you operate on data.

Check out the `% Change` column in `compensation`.  It shows the percentage increase in the CEO's total pay from the previous year.  For CEOs with no previous year on record, it instead says "(No previous year)".  The values in this column are *strings*, not numbers, so like the `Cash Pay` and `Total Pay` columns, it's not usable without a bit of extra work.

Given your current pay and the percentage increase from the previous year, you can compute your previous year's pay.  For example, if your pay is $\$120$ this year, and that's an increase of 50% from the previous year, then your previous year's pay was $\frac{\$120}{1 + \frac{50}{100}}$, or \$80.

**Question 3.5** Create a new table called `with_previous_compensation`.  It should be a copy of `cash_compensation`, but with the "(No previous year)" CEOs filtered out, and with an extra column called `2014 Total Pay ($)`.  That column should have each CEO's pay in 2014.

*Hint 1:* You can print out your results after each step to make sure you're on the right track.

*Hint 2:* We've provided a structure that you can use to get to the answer. However, if it's confusing, feel free to delete the current structure and approach the problem your own way!


```python
# Definition to turn percent to number
def percent_string_to_num(percent_string):
    """Converts a percentage string to a number."""
    return ...

# Compensation table where there is a previous year
having_previous_year = ...

# Get the percent changes as numbers instead of strings
# We're still working off the table having_previous_year
percent_changes = ...

# Calculate the previous year's pay
# We're still working off the table having_previous_year
previous_pay = ...

# Put the previous pay column into the having_previous_year table
with_previous_compensation = ...

with_previous_compensation
```

```python
grader.check("q35")
```

**Question 3.6** Determine the aggregate total pay in 2014 of all the CEOs that appear in the `with_previous_compensation` table. Assign this value to the variable `total_pay_all_ceos_2014`.


```python
total_pay_all_ceos_2014 = ...
total_pay_all_ceos_2014
```

```python
grader.check("q36")
```

## 4. Histograms
Earlier, we computed the average cash pay among the CEOs in our 102-CEO dataset.  The average doesn't tell us everything about the amounts CEOs are paid, though.  Maybe just a few CEOs make the bulk of the money, even among these 102.

We can use a *histogram* method to display the *distribution* of a set of numbers.  The table method `hist` takes a single argument, the name of a column of numbers.  It produces a histogram of the numbers in that column.

**Question 4.1.** Make a histogram of the cash pay of the CEOs in `cash_compensation`. Check with a peer or instructor to make sure you have the right plot. *If you get a warning, ignore it.*

*Hint:* If you aren't sure how to create a histogram, refer to the [Python Reference sheet](http://data8.org/fa21/python-reference.html).

```python
...
```

**Question 4.2.** How many CEOs made at least $10 million in cash pay? Find the value using code, then check that the value you found is consistent with what you see in the histogram.

*Hint:* Use the table method `where` and the property `num_rows`.


```python
num_ceos_at_least_10_million = ...
num_ceos_at_least_10_million
```

```python
grader.check("q42")
```

<img src="bolt.jpg" alt="drawing" width="300"/>

Bolt wants to congratulate you on finishing lab 4! You can now relax like Bolt!

## 6. Submission

Great job! You're finished with Lab 4!

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q11 (points: None)
```
>>> convert_to_percentage(0.35) == 35.0
True
```
```
>>> thirty_percent == 30.0
True
```


Tests for q13 (points: None)
```
>>> num_string_characters('Go Laney!') == 7
True
```


Tests for q14 (points: None)
```
>>> print_kth_top_movie_name(4)
The movie at rank 4 for total gross movie sales was: Avatar
```


Tests for q21 (points: None)
```
>>> import numpy
>>> cash_pay_type != numpy.ndarray
True
```
```
>>> 'str' in str(cash_pay_type)
True
```
```
>>> isinstance(cash_pay_type, type)
True
```
```
>>> cash_pay_type != int
True
```


Tests for q22 (points: None)
```
>>> str(marissa_mayer_pay_string)
'$1.00 '
```


Tests for q23 (points: None)
```
>>> type(marissa_mayer_pay) != str
True
```
```
>>> marissa_mayer_pay != 1
True
```
```
>>> marissa_mayer_pay == 1000000
True
```


Tests for q24 (points: None)
```
>>> convert_pay_string_to_number('$100 ') == 100000000.0
True
```
```
>>> convert_pay_string_to_number('$23 ') == 23000000.0
True
```


Tests for q31 (points: None)
```
>>> len(four_functions)
4
```
```
>>> callable(four_functions.item(0))
True
```
```
>>> callable(four_functions.item(1))
True
```
```
>>> callable(four_functions.item(2))
True
```
```
>>> callable(four_functions.item(3))
True
```


Tests for q32 (points: None)
```
>>> 'Cash Pay ($)' in cash_compensation.labels
True
```
```
>>> t = cash_compensation.sort('Cash Pay ($)', descending=True)
>>> np.isclose(t.column('Cash Pay ($)').item(0), 24890000.0)
True
```


Tests for q33 (points: None)
```
>>> import math
>>> math.isclose(average_cash_pay, 3208217.8217821782, rel_tol=0.1)
True
```


Tests for q34 (points: None)
```
>>> len(equity_proportion) == 101
True
```
```
>>> import math
>>> math.isclose(equity_proportion.item(3), 0.82352941, rel_tol=0.001)
True
```


Tests for q35 (points: None)
```
>>> '(No previous year)' not in with_previous_compensation.column('% Change')
True
```
```
>>> import math
>>> t = with_previous_compensation.sort('2014 Total Pay ($)', descending=True)
>>> value = t.column('2014 Total Pay ($)').item(0)
>>> math.isclose(value, 67700000.0, rel_tol=0.01)
True
```
```
>>> with_previous_compensation.num_rows == 80
True
```


Tests for q36 (points: None)
```
>>> import math
>>> math.isclose(total_pay_all_ceos_2014, 943583265.3638781, rel_tol=0.01)
True
```


Tests for q42 (points: None)
```
>>> num_ceos_at_least_10_million == 3
True
```
