```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

# Homework 3: Table Manipulation and Visualization

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Reading**: 
* [Visualization](https://inferentialthinking.com/chapters/07/Visualization.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)

# These lines do some fancy plotting magic.\n",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
```

## 1. Unemployment

The Great Recession of 2008-2009 was a period of economic decline observed globally, with scale and timing varying from country to country. In the United States, it resulted in a rapid rise in unemployment that affected industries and population groups to different extents.

The Employment Development Department publishes data about jobs in CA.  Below, we've loaded data on unemployment in the CA. The unemployment rate represents the number of unemployed as a percentage of the labor force. Labor force data are restricted to people 16 years of age and older, who currently reside in CA, who do not reside in institutions (e.g., penal and mental facilities, homes for the aged), and who are not on active duty in the Armed Forces.

Over the course of a year, the size of the labor force, the levels of employment and unemployment, and other measures of labor market activity undergo fluctuations due to seasonal events including changes in weather, harvests, major holidays, and school schedules. Because these seasonal events follow a more or less regular pattern each year, their influence on statistical trends can be eliminated by seasonally adjusting the statistics from month to month. These seasonal adjustments make it easier to observe the cyclical, underlying trend, and other nonseasonal movements in the series. The data provided includes seasonally adjusted and non-seasonally adjusted data.

1. Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. Among people who are able to work and are looking for a full-time job, the percentage who can't find any job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", so the acronym for this index is NEI-PTER.  (Economists are great at marketing.)

The source of the data is [here](https://fred.stlouisfed.org/categories/33509).

**Question 1.** The data are in a CSV file called `Unemployment_rate_CA.csv`.  Load that file into a table called `ca_unemployment`. **(4 Points)**


```python
ca_unemployment = ...
ca_unemployment
```

```python
grader.check("q1_1")
```

**Question 2.** Sort the data in descending order by the Seasonally Adjusted rate, naming the sorted table `by_season`.  Create another table called `by_not_season` that's sorted in descending order by the Not Seasonally Adjusted rate instead. **(4 Points)**


```python
by_season = ...
by_not_season = ...
```

```python
grader.check("q1_2")
```

```python
# Run this cell to check your by_season table. You do not need to change the code.
by_season.show(5)
```

```python
# Run this cell to check your by_not_season table. You do not need to change the code.
by_not_season.show(5)
```

**Question 3.** Use `take` to make a table containing the data for the 12 months when the Seasonally Adjusted unemployment rate was smallest.  Call that table `smallest_season`.

`smallest_season` should be sorted in ascending order of `Seasonally Adjusted`. Note that each row of `ca_unemployment` represents a month. **(4 Points)**


```python
smallest_season = ...
smallest_season
```

```python
grader.check("q1_3")
```

**Question 4.** The unemployment rates differ depending on whether you look at the Seasonally Adjusted or Not Seasonally Adjusted rates. Typically, the monthly employment and unemployment numbers reported in the news are seasonally adjusted data. Seasonally adjusted data are useful when comparing several months of data. Annual average estimates are calculated from the not seasonally adjusted data series.

Compute an array containing the difference between the Seasonally Adjusted and Not Seasonally Adjusted rates.  (The first element of the array should correspond to the first row of `ca_unemployment`, and so on.) **(4 Points)**

*Note:* Use the original `ca_unemployment` table for this.


```python
diff = ...
diff
```

```python
grader.check("q1_4")
```

**Question 5.** Add `diff` as a column to `ca_unemployment` (name the column `Difference`) and sort the resulting table by that column in descending order.  Call the resulting table `by_diff`.

Try to do this with a single line of code, if you can. **(4 Points)**


```python
by_diff = ...
by_diff
```

```python
grader.check("q1_5")
```

**Question 6.** Create a line plot of Seasonally Adjusted rates over time. To do this, create a new table called `season_over_time` by making a copy of the `ca_unemployment` table and including only the columns `Date` and `Seasonally Adjusted`, and adding a new column: `Year` using the `year` array. Then, generate a line plot using one of the table methods you've learned in class.

The order of the columns matter for our correctness tests, so be sure `Year` comes before `Seasonally Adjusted`. **(4 Points)**

*Note:* When constructing `season_over_time`, do not just add the `year` column to the `by_season` table. Please follow the directions in the question above.


```python
year = 1976 + np.arange(by_season.num_rows)/10
year
season = ca_unemployment.column("Seasonally Adjusted")
season_over_time = ...
season_over_time
...
plots.ylim(0,20); # Do not change this line
```

```python
grader.check("q1_6")
```

**Question 7.** Were Seasonally Adjusted unemployment rates higher during the Great Recession (that is in the years 2008 through 2011) than during 2020? Assign `high_rates` to `True` if you think Seasonally Adjusted rates were higher during the Great Recession time period, or `False` if you think they weren't. **(4 Points)**


```python
high_rates = ...
```

```python
grader.check("q1_7")
```

## 2. Birth Rates

The following table gives Census-based population estimates for each US state on both July 1, 2015 and July 1, 2016. The last four columns describe the components of the estimated change in population during this time interval. **For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico and the District of Columbia.**

The data was taken from [here](http://www2.census.gov/programs-surveys/popest/datasets/2010-2016/national/totals/nst-est2016-alldata.csv). (Note: If it doesn't download for you when you click the link, please copy and paste it into your address bar!) If you want to read more about the different column descriptions, click [here](http://www2.census.gov/programs-surveys/popest/datasets/2010-2015/national/totals/nst-est2015-alldata.pdf).

The raw data is a bit messy—run the cell below to clean the table and make it easier to work with.

```python
# Don't change this cell; just run it.
pop = Table.read_table('nst-est2016-alldata.csv').where('SUMLEV', 40).select([1, 4, 12, 13, 27, 34, 62, 69])
pop = pop.relabeled('POPESTIMATE2015', '2015').relabeled('POPESTIMATE2016', '2016')
pop = pop.relabeled('BIRTHS2016', 'BIRTHS').relabeled('DEATHS2016', 'DEATHS')
pop = pop.relabeled('NETMIG2016', 'MIGRATION').relabeled('RESIDUAL2016', 'OTHER')
pop = pop.with_columns("REGION", np.array([int(region) if region != "X" else 0 for region in pop.column("REGION")]))
pop.set_format([2, 3, 4, 5, 6, 7], NumberFormatter(decimals=0)).show(5)
```

**Question 1.** Assign `west_death_rate` to the total annual death rate of the western US during this time interval. The annual death rate for a year-long period is the total number of deaths in that period as a proportion of the total population size at the start of the time period. **(4 Points)**

*Hint:* Which year corresponds to the start of the time period?

*Hint:* Which `REGION` in the table corresponds to the states in the western US?



```python
west_region = pop.where("REGION",4)
west_death_rate = ...
west_death_rate
```

```python
grader.check("q2_1")
```

**Question 2.** Assign `increase` to the number of western US states for which the **annual number of Natural Increase** was lower than 15,000. **Natural Increase** is the difference between the numbers of births and deaths in a population. The `BIRTHS` column contains the number of net birth counts by state. The `DEATHS` column contains the number of net death counts by state.**(4 Points)**

*Hint*: `west_ni_number` should be a table and `increase` should be a number.

*Hint*: use name that you for `REGION` 4 in previous problem.


```python
west_ni_number = ...
increase = ...
increase

```

```python
grader.check("q2_2")
```

**Question 3.** Assign `west_increase` to an array that returns the names of all western US states in the order of highest to lowest **Natural Increase.** **(4 Points)**

*Hint:* use parts of previous answer.


```python
west_increase = ...
west_increase
```

```python
grader.check("q2_3")
```

**Question 4.** In the next question, you will be creating a visualization to understand the relationship between western states and **Natural Increase**. 

What visualization is most appropriate to determine which western state has the highest **Natural Increase?**

1. Line Graph
2. Bar Chart
3. Scatter Plot

Assign `visualization` below to the number corresponding to the correct visualization. **(4 Points)**


```python
visualization = ...
```

```python
grader.check("q2_4")
```

<!-- BEGIN QUESTION -->

**Question 5.** In the code cell below, create a visualization that will help us determine if there is an association between western US birth rate and death rate during this time interval. It may be helpful to create an intermediate table here. **(4 Points)**

Things to consider:

- What type of chart will help us illustrate an association between 2 variables?
- How can you manipulate a certain table to help generate your chart?
- Check out the Recommended Reading for this homework!


```python
# In this cell, use birth_rates and death_rates to generate your visualization
birth_rates = west_region.column('BIRTHS') / west_region.column('2015')
death_rates = west_region.column('DEATHS') / west_region.column('2015')
...
```

<!-- END QUESTION -->

**Question 6.** True or False: There is an association between western US birth rate and death rate during this time interval. 

Assign `assoc` to `True` or `False` in the cell below. **(4 Points)**


```python
assoc = ...
```

```python
grader.check("q2_6")
```

## 3. Flint, Michigan

Flint, Michigan is a city about an hour outside Detroit. In 2014, Flint, Michigan switched its drinking water supply from Detroit’s system to the Flint River in a cost-saving move.  In September 2015 a Flint pediatrician found: The incidence of elevated blood-lead levels in children citywide had nearly doubled since 2014—and nearly tripled in certain neighborhoods. In 2015, an independent group of researchers at Virginia Tech, volunteered to conduct a study into the water crisis. They collected Water samples from resident homes throughout the city and measured the lead levels. Below we load a table containing water samples collected from 272 homes from the [Flint Water Study](https://flintwaterstudy.org/2015/12/complete-dataset-lead-results-in-tap-water-for-271-flint-samples/). 

The `Sampleid` is the ID assigned to each household, and `Zip Code` is the zip code of the household. The `PB Bottle 1 (ppb) - First Draw` is the amount of lead (Pb) measured in parts per billion (ppb) in the water as soon as the tap turned on. The `PB Bottle 2 (ppb) - 45 secs flushing` is the amount of lead (Pb) measured in parts per billion (ppb) after water from the tap has been flushing for 45 seconds. The `PB Bottle 3 (ppb) - 2 min flushing` is the amount of lead (Pb) measured in parts per billion (ppb) after water from the tap has been flushing for 2 minutes. 

```python
flint = Table.read_table("Flint.csv")
print("Flint, Michigan")
flint.show(4)
```

<!-- BEGIN QUESTION -->

**Question 1.** Produce a histogram that visualizes the distributions of lead (Pb) measured in parts per billion (ppb) after water from the tap has been flushing for 2 minutes using the given bins in `equal_bins`. **(4 Points)**


```python
equal_bins = np.arange(0, 20, 1.5)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.** Now, produce a histogram that visualizes the distribution of lead (Pb) measured in parts per billion (ppb) after water from the tap has been flushing for 45 seconds using the given bins in `equal_bins`. **(4 Points)**


```python
equal_bins = np.arange(0, 20, 1.5)
...

# Don't delete the following line!
plots.ylim(0, .5);
```

<!-- END QUESTION -->

**Question 3.** Let's take a closer look at the y-axis label. Assign `unit_meaning` to an integer (1, 2, 3) that corresponds to the "unit" in "Percent per unit". **(4 Points)**

1. second  
2. minute
3. ppb (parts per billion of lead Pb)


```python
unit_meaning = ...
unit_meaning
```

```python
grader.check("q3_3")
```

**Question 4.** Assign `flush_45sec_under_4_point_5` to the percentage of bottles that contain less than 4.5 ppb of lead after 45 seconds of flushing. Assign `flush_2min_under_4_point_5` to the percentage of bottles that contain less than 4.5 ppb of lead after 2 minutes of flushing. Use the height variables provided below in order to compute the percentages. Your solution should only use height variables, numbers, and mathematical operations. You should **not** access the table `Flint` in any way. **(4 Points)**


```python
flush_45sec_under__1_point_5_height = 34.9
flush_2min_under_1_point_5_height = 41.8
flush_45sec_1_point_5_to_under_3_height = 9.3
flush_2min_1_point_5_to_under_3_height = 9.1
flush_45sec_3_to_under_4_point_5_height = 5.9
flush_2min_3_to_under_4_point_5_height = 5.4

flush_45_under_4_point_5 = ...
flush_2min_under_4_point_5 = ...

flush_45_under_4_point_5, flush_2min_under_4_point_5
```

```python
grader.check("q3_4")
```

**Question 5.** Let's take a closer look at the distribution of lead afer 45 seconds of flushing. Assign `flush_2min_median_bin` to an integer (1, 2, 3, or 4) that corresponds to the bin that contains the median amout of lead. **(4 Points)**

1. 0-1.5 ppb  
2. 1.5-3 ppb  
3. 3-4.5 ppb  
4. 4.5-6 ppb  

*Hint:* The median of a sorted list has half of the list elements to its left, and half to its right.


```python
flush_2min_median_bin = ...
flush_2min_median_bin
```

```python
grader.check("q3_5")
```

<!-- BEGIN QUESTION -->

**Question 6.** Identify one difference between the histograms, in terms of the statistical properties. Can you comment on the average and/or skew of each histogram? **(4 Points)**

*Hint*: The best way to do this is to compare the two histograms (from 3.1 and 3.2) visually.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7.** 
A. Does flushing the tap for 2 minutes decrease the amount of lead in the tap water compared to flushing the tap for 45 seconds? Explain your answer using at least one complete sentence.

B. According to the Virginia Tech researchers, lead levels of over 5 parts per billion (ppb) indicate there is a “very serious” problem in the water supply. If lead levels are over 15 ppb, the federal law requires that corrective action be taken. Based on this information, does the data you’ve analyzed suggest that there was a serious problem of lead in the water supply in Flint, Michigan in 2015? Explain your answer using at least 1-2 complete sentences. Use numerical data (with units) to support your answer.

**(4 Points)**



_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Histograms

Consider the following scatter plot: 

![](scatter.png)

The axes of the plot represent values of two variables: $x$ and $y$. 

Suppose we have a table called `t` that has two columns in it:

- `x`: a column containing the x-values of the points in the scatter plot
- `y`: a column containing the y-values of the points in the scatter plot

Below, you are given three histograms—one corresponds to column `x`, one corresponds to column `y`, and one does not correspond to either column. 

**Histogram A:**

![](var2.png) 

**Histogram B:**

![](var3.png)

**Histogram C:**

![](var1.png)

**Question 1.** Suppose we run `t.hist('x')`. Which histogram does this code produce? Assign `histogram_column_x` to either 1, 2, or 3. **(5 Points)**

1. Histogram A
2. Histogram B
3. Histogram C


```python
histogram_column_x = ...
```

```python
grader.check("q4_1")
```

<!-- BEGIN QUESTION -->

**Question 2.** State at least one reason why you chose the histogram from Question 1. **Make sure to clearly indicate which histogram you selected** (ex: "I chose histogram A because ..."). **(5 Points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.** Suppose we run `t.hist('y')`. Which histogram does this code produce? Assign `histogram_column_y` to either 1, 2, or 3. **(5 Points)**

1. Histogram A
2. Histogram B
3. Histogram C


```python
histogram_column_y = ...
```

```python
grader.check("q4_3")
```

<!-- BEGIN QUESTION -->

**Question 4.** State at least one reason why you chose the histogram from Question 3.  **Make sure to clearly indicate which histogram you selected** (ex: "I chose histogram A because ..."). **(5 Points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

You're done with Homework 3!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [4])
```
>>> ca_unemployment.select('Date', 'Seasonally Adjusted', 'Not Seasonally Adjusted').take(0)
Date     | Seasonally Adjusted | Not Seasonally Adjusted
1/1/1976 | 9.2                 | 10.4
```


Tests for q1_2 (points: [2, 2])
```
>>> by_season.take(100)
Area Name  | Date      | Year | Month   | Seasonally Adjusted | Not Seasonally Adjusted
California | 10/1/2020 | 2020 | October | 9.3                 | 8.8
```
```
>>> by_not_season.take(100)
Area Name  | Date     | Year | Month | Seasonally Adjusted | Not Seasonally Adjusted
California | 3/1/1977 | 1977 | March | 8.9                 | 9.3
```


Tests for q1_3 (points: [2])
```
>>> smallest_season.take(0)
Area Name  | Date     | Year | Month | Seasonally Adjusted | Not Seasonally Adjusted
California | 7/1/2022 | 2022 | July  | 3.8                 | 4.1
```


Tests for q1_4 (points: [1, 3])
```
>>> round(diff.item(6), 4) == -0.5
True
```
```
>>> round(diff.item(6), 4)
-0.5
```


Tests for q1_5 (points: [4])
```
>>> by_diff.take(0)
Area Name  | Date     | Year | Month | Seasonally Adjusted | Not Seasonally Adjusted | Difference
California | 5/1/1976 | 1976 | May   | 9                   | 7.9                     | 1.1
```


Tests for q1_6 (points: [0])
```
>>> season_over_time.take(0)
Date     | Year | Seasonally Adjusted
1/1/1976 | 1976 | 9.2
```


Tests for q1_7 (points: [0])
```
>>> high_rates == True or high_rates == False
True
```


Tests for q2_1 (points: [0])
```
>>> 0 < west_death_rate < 1
True
```


Tests for q2_2 (points: [0])
```
>>> 0 < increase <= 10
True
```


Tests for q2_3 (points: [0])
```
>>> west_increase == west_increase
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True,  True,  True], dtype=bool)
```


Tests for q2_4 (points: [0, 0])
```
>>> type(visualization) == int
True
```
```
>>> 1 <= visualization <= 3
True
```


Tests for q2_6 (points: [0])
```
>>> type(assoc) is bool
True
```


Tests for q3_3 (points: [0])
```
>>> 1 <= unit_meaning <= 3
True
```


Tests for q3_4 (points: [0, 0])
```
>>> flush_45_under_4_point_5 >= 0 and flush_45_under_4_point_5 <= 100
True
```
```
>>> flush_2min_under_4_point_5 >= 0 and flush_2min_under_4_point_5 <= 100
True
```


Tests for q3_5 (points: [0])
```
>>> 1 <= flush_2min_median_bin <= 4
True
```


Tests for q4_1 (points: [0, 1])
```
>>> type(histogram_column_x) == int
True
```
```
>>> histogram_column_x == 1 or histogram_column_x == 2 or histogram_column_x == 3
True
```


Tests for q4_3 (points: [0, 1])
```
>>> type(histogram_column_y) == int
True
```
```
>>> histogram_column_y == 1 or histogram_column_y == 2 or histogram_column_y == 3
True
```
