```python
# Initialize Otter
import otter
grader = otter.Notebook("hw09.ipynb")
```

# Homework 9: Sample Sizes and Confidence Intervals

**Helpful Resource:**

- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 

* [Estimation](https://inferentialthinking.com/chapters/13/Estimation.html)
* [Why the Mean Matters](https://inferentialthinking.com/chapters/14/Why_the_Mean_Matters.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1.  Understanding Distributions in NBA Player Careers

Recent statistics indicate that the average NBA career lasts 5 years with a standard deviation of 1.5 years. This statistical insight allows us to explore and understand the distribution of players' careers within the league, applying Chebyshev's inequality.

In each part below, fill in the blank with a percent that makes the statement true **without further assumptions**, and explain your answer.

*Note:* No credit will be given for loose bounds such as "at least 0%" or "at most 100%". Give the best answer that is possible with the information given.

<!-- BEGIN QUESTION -->

**Question 1.1.**  (Years in the NBA): At least _______% of NBA players have careers that last between 2 and 8 years in the league. **(6 Points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.2.** (Outside Career Range): At most _______% of NBA players have careers that do not last between 2 and 8 years in the league.

 **(6 Points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.3.** (Long Careers): At most _______% of NBA players have careers that last more than 8 years in the league. **(6 Points)**

*Hint:* If you're stuck, try thinking about what the distribution may look like in this case.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Sample Size and Confidence Intervals in Basketball Analytics

A data science class at a university, renowned for its basketball team, wants to estimate the percentage of students at the school who created a bracket for the NCAA March Madness college basketball tournament. To do this, they plan to take a random sample of students. You can assume their method of sampling is equivalent to drawing at random with replacement from the student body.

***Please review [Section 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html#) of the textbook before proceeding with this section. Your work will go much faster that way!***

**Question 2.1.**  Assign `fewest` to the smallest number of students they should sample to ensure that an **95%** confidence interval for the percentage of NBA followers has a width of no more than 4% from left end to right end. **(6 points)**

*Hint:* How can our data be represented to show if a student in the sample follows the NBA? Given this, what assumptions can we make for the SD of the population? [Section 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html#) might be helpful!

*Note:* While the true smallest sample size would have to be an integer, please leave your answer in decimal format for the sake of our tests.


```python
fewest = ...
fewest
```

```python
grader.check("q2_1")
```

<!-- BEGIN QUESTION -->

**Question 2.2.** Suppose the data science class decides to construct a 99% confidence interval instead of a 95% confidence interval, but they still require that the width of the interval is no more than 4% from left end to right end. Will they need the same sample size as in 2.1? Pick the right answer and explain further without calculation. **(6 Points)**

1. Yes, they must use the same sample size.
2. No, a smaller sample size will work.
3. No, they will need a bigger sample.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.3:** The professor explains that a 99% confidence interval for the parameter is constructed similarly to a 95% confidence interval, except that you have to go 2.58 SDs on either side of the estimate (± 2.58) instead of 2 SDs on either side (±2). Assign `fewest_num` to the smallest number of students they should sample to ensure that a **99%** confidence interval for the percentage of March Madness followers has a width of no more than 4% from left end to right end. **(6 points)**

*Note:* While the true smallest sample size would have to be an integer, please leave your answer in decimal format for the sake of our tests.


```python
fewest_num = ...
fewest_num
```

```python
grader.check("q2_3")
```

For this next exercise, please consult [Section 14.3.4](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html#the-standard-normal-cdf) of the textbook for similar examples.

The students are curious about how the professor came up with the value 2.58 in Question 2.3. She says she ran the following two code cells. The first one calls the `datascience` library function `plot_normal_cdf`, which displays the proportion that is at most the specified number of SDs above average under the normal curve plotted with standard units on the horizontal axis. You can find the documentation [here](http://data8.org/datascience/util.html#datascience.util.plot_normal_cdf).

*Note:* The acronym `cdf` stands for `cumulative distribution function`. It measures the proportion to the left of a specified point under a probability histogram.

```python
plot_normal_cdf(2.58)
```

To run the second cell, the professor had to first import a Python library for probability and statistics:

```python
# Just run this cell
from scipy import stats
```

Then she used the `norm.cdf` method in the library to find the gold proportion above.

```python
# Just run this cell
stats.norm.cdf(2.58)
```

*Note*: You do not need to understand how the `scipy` library or how to use the method yourself.

<!-- BEGIN QUESTION -->

**Question 2.4:** The professor used a function to show that the percentage in a normal distribution that is at most 2.58 SDs above average is about **99%.** Explain why 2.58 is the right number of SDs to use when constructing a **99%** confidence interval. **(6 Points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Just run this cell, do not change it.
stats.norm.cdf(1.285)
```

**Question 2.5.** The cell above shows that the proportion that is at most 1.285 SDs above average in a normal distribution is 90%. 
If you start at the estimate and go 1.285 SDs on either side, then you will get a _______% confidence interval for the parameter. **(6 points)**

1. 80
2. 90
3. 95
4. 99



```python
option = ...
option
```

```python
grader.check("q2_5")
```

## 3. Polling and the Normal Distribution

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

The university is facing budget cuts. The university’s admistration is considering a new initiative, Proposition Steal, which would significantly decrease funds allocated to student recreational facilities, particularly basketball courts. They want to gauge student **opposition** for the initiative. If more than 70% of students oppose Proposition Steal, it could set off campus protests and demonstrations.

Moesha is a data science consultant hired by the administration. To estimate student opposition, she takes a uniform random sample of university students. In her survey, she finds 420 out of 625 students sampled will oppose Proposition Steal. We have provided a table for you below which has 3 columns: the first two columns are identical to `sample`. The first column indicates support or opposition, the second column indicates the count of students. The third column contains the proportion of the total sampled students that chose each option.

```python
sample = Table().with_columns(
    "Opinion",  make_array("Support", "Oppose"),
    "Count", make_array(205,   420))

sample_size = sum(sample.column("Count"))
sample_with_proportions = sample.with_column("Proportion", sample.column("Count") / sample_size)
sample_with_proportions
```

**Question 3.1.** Moesha wants to use 7,000 bootstrap resamples to compute a confidence interval for the proportion of university students who will oppose Proposition Steal.

Using the data from the survey where 420 out of 625 students opposed the initiative, she will create an empirical distribution of support proportions. This process involves using bootstrap resampling to simulate 7,000 survey outcomes, which will help estimate the variability of the sample proportion. Assign `resample_support_proportions` to contain the Oppose proportion of each bootstrap resample. Then, visualize `resample_support_proportions` with a histogram.
**(6 points)**

*Hint:* `sample_proportions` may be useful here!
*Hint* your histogram should have a bell curve shape.

```python
resample_oppose_proportions = make_array()
for i in np.arange(7000):
    resample = ...
    resample_oppose_proportions = ...
Table().with_column("Resample Oppose proportion", resample_oppose_proportions).hist(bins=np.arange(.2, .8, .01))
```

```python
grader.check("q3_1")
```

<!-- BEGIN QUESTION -->

**Question 3.2.** Explain how the Central Limit Theorem (CLT) is relevant to the situation at hand and describe its influence on the observed distribution in the histogram from the bootstrap resamples. **(6 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

In a population whose members are 0 and 1, there is a simple formula for the **standard deviation of that population**:

$$\text{standard deviation of population} = \sqrt{(\text{proportion of 0s}) \times (\text{proportion of 1s})}$$

(Figuring out this formula, starting from the definition of the standard deviation, is an fun exercise for those who enjoy algebra.)

**Question 3.3.** Using only the Central Limit Theorem and the numbers of Support and Oppose opinions in our sample of 625, *algebraically* compute the predicted standard deviation of the `resample_support_proportions` array. Assign this number to `approximate_sd`. **Do not access the data in `resample_not_support_proportions` in any way.** **(6 points)**

Remember that the standard deviation of the sample means can be computed from the population SD and the size of the sample (the formula above might be helpful). If we do not know the population SD, we can use the sample SD as a reasonable approximation in its place. [This section](https://inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html#the-sd-of-all-the-sample-means) of the textbook also may be helpful.


```python
approx_pop_sd = ...
approximate_sd = ...
approximate_sd
```

```python
grader.check("q3_3")
```

**Question 3.4.** Calculate the standard deviation of the `resample_oppose_proportions` array. This value serves as an estimated standard deviation of the potential sample proportions and can be used to check the accuracy of your response to the subsequent question **(6 points)**


```python
exact_sd = ...
exact_sd
```

```python
grader.check("q3_4")
```

**Question 3.5.** **Without referencing the resample_oppose_proportions dataset,** calculate an estimated 95% confidence interval for the proportion of students who oppose Proposition Steal. **(6 points)**

The following visualization will display your calculated interval as a red bar beneath the histogram of bootstrap proportions to assist in confirming the plausibility of your calculated interval.


*Hint:* How many SDs corresponds to 95% of the distribution promised by the CLT? Recall the discussion in the textbook [here](https://inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html).

*Hint:* The `approximate_sd` that you computed earlier could be useful in this step!


```python
lower_limit = ...
upper_limit = ...
print('lower:', lower_limit, 'upper:', upper_limit)
```

```python
grader.check("q3_5")
```

```python
# Run this cell to plot your confidence interval.
Table().with_column("Resample Oppose proportion", resample_oppose_proportions).hist(bins=np.arange(.5, .8, .01))
plt.plot(make_array(lower_limit, upper_limit), make_array(0, 0), c='r', lw=10);
```

<div class="hide">\pagebreak</div>
Your confidence interval should include the proportion 0.7. This means we can't be sure less than 70% of students will oppose Proposition Steal, even though the sample oppose proportion is less than 70%.

The administration needs to know whether whether opposition will be below 70%. Absolute certainty would require surveying the entire student body, which is impractical. However, they would be satisfied if the standard deviation of the sample proportion were reduced to 0.007. They ask Moesha to conduct a new survey with a sample size large enough to achieve that. Since conducting the new survey is expensive, they ask the sample not be any larger than necessary.

Moesha turns to Chapter 14 of the statistics textbook. Instead of the cautious approach that assumes the maximum variability for a proportion (using 0.5 to represent the spread between support and non-support as binary outcomes), she opts to base her calculations on the observed standard deviation from the sample data.

$$\sqrt{(\text{Supported proportion in the sample}) \times (\text{Not supported proportion in the sample})}.$$

Under that assumption, Moesha decides that a sample of 4,460 students would suffice.

**Question 3.6.** Does Moesha's sample size achieve the desired standard deviation of sample means? What SD would you achieve with a smaller sample size? A higher sample size?

To explore this, first compute the SD of sample means obtained by using Moesha's sample size. **(6 points)**


```python
estimated_population_sd = ...
moesha_sample_size = ...
moesha_sample_mean_sd = ...
print("With Moesha's sample size, you would predict a sample mean SD of %f." % moesha_sample_mean_sd)
```

```python
grader.check("q3_6")
```

**Question 3.7.** Now, calculate the standard deviation of the sample means that would result from using a smaller sample size. Choose a sample size that is notably less than what Moesha used, though any size smaller than hers is acceptable. **(5 points)**


```python
smaller_sample_size = ...
smaller_sample_mean_sd = ...
print("With this smaller sample size, you would predict a sample mean SD of %f" % smaller_sample_mean_sd)
```

```python
grader.check("q3_7")
```

**Question 3.8.** Next, determine the standard deviation of the sample means expected from using a larger sample size. Opting for a significantly larger number will more clearly illustrate the differences, but any size greater than what Moesha utilized will suffice.. **(5 points)**



```python
larger_sample_size = ...
larger_sample_mean_sd = ...
print("With this larger sample size, you would predict a sample mean SD of %f" % larger_sample_mean_sd)
```

```python
grader.check("q3_8")
```

**Question 3.9.** Considering this analysis, evaluate whether Moesha's chosen sample size was essentially the smallest adequate sample, assuming the sample standard deviation matches the population standard deviation. Set `min_sufficient` to `True` if 4460 was close to the smallest adequate sample size, and to `False` otherwise. **(4 points)**

Considering this analysis, evaluate whether Moesha's chosen sample size was essentially the smallest adequate sample, assuming the sample standard deviation matches the population standard deviation. Set `min_sufficient` to `True` if 8,665 was close to the smallest adequate sample size, and to `False` otherwise


```python
min_sufficient = ...
min_sufficient
```

```python
grader.check("q3_9")
```

You're done with Homework 9!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q2_1 (points: [0])
```
>>> 1100 <= fewest <= 2550
True
```


Tests for q2_3 (points: [0])
```
>>> 4000 <= fewest_num <= 4200
True
```


Tests for q2_5 (points: [0])
```
>>> type(option) == int
True
```


Tests for q3_1 (points: [1, 5])
```
>>> len(resample_oppose_proportions) == 7000
True
```
```
>>> abs(np.mean(resample_oppose_proportions) - 0.648) < 0.025 and np.std(resample_oppose_proportions) < 0.02
True
```


Tests for q3_3 (points: [0])
```
>>> approximate_sd < 0.02
True
```


Tests for q3_4 (points: [0])
```
>>> exact_sd < 0.02
True
```


Tests for q3_5 (points: [0])
```
>>> 0.6 <= lower_limit < upper_limit <= 0.71
True
```


Tests for q3_6 (points: [0])
```
>>> type(moesha_sample_size) == int
True
```


Tests for q3_7 (points: [1])
```
>>> smaller_sample_size < moesha_sample_size
True
```


Tests for q3_8 (points: [1])
```
>>> larger_sample_size > moesha_sample_size
True
```


Tests for q3_9 (points: [0])
```
>>> min_sufficient or not min_sufficient
True
```
