```python
# Initialize Otter
import otter
grader = otter.Notebook("hw08.ipynb")
```

# Homework 8: Confidence Intervals

**Helpful Resource:**

- [Python Reference](http://data8.org/fa23/reference/): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Reading**: 

* [Estimation](https://www.inferentialthinking.com/chapters/13/Estimation)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore')
```

## 1. Athletic Shoe Brands Preferred by Laney College Students

Camron and Jamal are curious to determine the most favored athletic shoe brand among Laney College students. They conduct a survey of 1,300 Laney College students, chosen uniformly at random, inquiring about each student's preferred athletic shoe brand. (Note: The data mentioned here is created for the purposes of this assignment.) The options given were [Nike](https://www.nike.com/?cp=90083564018_search_-nike-g-20300414621-150548537237-e-c---681090968314-kwd-19354280-9032064&gad_source=1&gclid=Cj0KCQjw2a6wBhCVARIsABPeH1vU0ZDJ9KF1XXhS8AnRw74m80_2ifsykmyNncn502TWUh1QdJ5eRbUaAlF6EALw_wcB&gclsrc=aw.ds), [Adidas](https://www.adidas.com/us/sale?af_channel=Search&af_click_lookback=30d&af_reengagement_window=30d&c=adidas-Trademark-General-B-new&cm_mmc=AdieSEM_Google-_-adidas-Trademark-General-B-new-_-Trademark-X-X-Unknown-Exact-_-43700014909128661&cm_mmca1=US&cm_mmca2=&ds_agid=58700001939460312&gad_source=1&gclid=Cj0KCQjw2a6wBhCVARIsABPeH1uWeEROEnKb3-1EnI0N1QxpwdnoxUnoUfBEuP8Gj7Rg5FaBlX5UlPcaAgb2EALw_wcB&gclsrc=aw.ds&is_retargeting=true&pid=googleadwords_temp), [Hoka](https://www.hoka.com/en/us/?utm_source=google&utm_medium=ppc_brand&gad_source=1&gclid=Cj0KCQjw2a6wBhCVARIsABPeH1vG6gJ2Ksc6RhSCklpZcpPm6n9ly49aE2-bCIh41VUdvq-XlRM5clsaAmxzEALw_wcB&gclsrc=aw.ds), and [Puma](https://us.puma.com/us/en?gad_source=1&gclid=Cj0KCQjw2a6wBhCVARIsABPeH1ufQhQ1gSVI1Ta1Qpo_2mFnsxf1jVnXBGfE2fMg3z45dHOlwA36Yh8aAkdDEALw_wcB). Upon analyzing the survey results, Camron and Jamal present the following percentages from their sample:

|Shoes  | Percentage|
|:------------:|:------------:|
|Nike | 44% |
|Adidas | 26% |
|Hoka | 17% |
|Puma | 13% |

These percentages are derived from a uniform random sample of Laney College's student population. Our goal is to estimate the corresponding *parameters*, or true preference rates for each shoe brand within the *population* (i.e., all students at Laney College) based on this sample data. We will use confidence intervals to compute a range of values that reflects the uncertainty of our estimates.


The table `preferences` contains the results of Camron and Jamal's survey.

```python
# Just run this cell
preferences = Table.read_table('preferences.csv')
preferences
```

**Question 1.1.** Complete the function `one_resampled_percentage` below. It should return Nike's ***percentage*** of times it was prefered after taking the original table (`tbl`) and performing one bootstrap sample of it. Reminder that a percentage is between 0 and 100. **(9 Points)**

*Note:* `tbl` will always be in the same format as `preferences`.

*Hint:* Given a table of votes, how can you figure out what percentage of the votes are for a certain shoe brand? **Be sure to use percentages, not proportions, for this question!**


```python
def one_resampled_percentage(tbl):
    ...

one_resampled_percentage(preferences)
```

```python
grader.check("q1_1")
```

**Question 1.2.** Complete the `percentages_in_resamples` function such that it simulates and returns an array of 2000 elements, where each element represents a bootstrapped estimate of the percentage of students who will prefer Nike. You should use the `one_resampled_percentage` function you wrote above. **(9 Points)**


```python
def percentages_in_resamples():
    percentage_nike = make_array()
    ...
```

```python
grader.check("q1_2")
```

In the following cell, we run the function you just defined, `percentages_in_resamples`, and create a histogram of the calculated statistic for the 2000 bootstrap estimates of the percentage of students who prefered Nike. 

*Note:* This might take a few seconds to run.

```python
resampled_percentages = percentages_in_resamples()
Table().with_column('Estimated Percentage', resampled_percentages).hist("Estimated Percentage")
```

**Question 1.3.** Using the array `resampled_percentages`, find the values at the two edges of the middle 90% of the bootstrapped percentage estimates. Compute the lower and upper ends of the interval, named `nike_lower_bound` and `nike_upper_bound` respectively. **(9 Points)**

*Hint:* If you are stuck on this question, try looking over [Chapter 13](https://inferentialthinking.com/chapters/13/Estimation.html) of the textbook.


```python
nike_lower_bound = ...
nike_upper_bound = ...
print(f"Bootstrapped 90% confidence interval for the percentage of students who prefered Nike in the population: [{nike_lower_bound:.2f}, {nike_upper_bound:.2f}]")
```

```python
grader.check("q1_3")
```

**Question 1.4.** The survey results seem to indicate that Nike is beating each of the other shoe brands among the students. However, all the other shoe brands combined have a lead over Nike. We would like to use confidence intervals to determine a range of likely values for the true lead all the other shoe brands combined have over Nike. The calculation for Nike's position behind Adidas, Hoka, and Puma combined is:

$$(\text{100 percent} - \text{Nike's percent of student preference})    -\text{Nike's percent of the student preference}$$

Define the function `one_resampled_difference` that returns from one bootstrap sample of `tbl`, **exactly one value** of the percentage lead Adidas, Hoka, and Puma combined have over Nike. **(9 Points)**

*Hint 1:* Nike's leading all the other brands individually but not combined. We need to create a confidence interval that reflects the combined lead of the other brands over Nike.

*Hint 2:* Given a table of preferences, how can you figure out what percentage of the students prefer a certain shoe brand? **Be sure to use percentages, not proportions, for this question!**

*Note:* If the skeleton code provided within the function is not helpful for you, feel free to approach the question using your own variables.


```python
def one_resampled_difference(tbl):
    bootstrap_sample = ...
    Nike_percentage = ...
    ...
```

```python
grader.check("q1_4")
```

<!-- BEGIN QUESTION -->

**Question 1.5.** Write a function called `leads_in_resamples` that returns an array of 2000 elements representing the bootstrapped estimates (the result of calling `one_resampled_difference`) of Adidas, Hoka, and Puma's combined lead over Nike. Afterwards, run the cell to plot a histogram of the resulting samples. **(9 Points)**

*Hint:* If you see an error involving `NoneType`, consider what components a function needs to have!


```python
def leads_in_resamples():
    ...

sampled_leads = leads_in_resamples()
Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")
```

<!-- END QUESTION -->

**Question 1.6.** Use the simulated data in `sampled_leads` from Question 1.5 to compute an approximate 90% confidence interval for Adidas, Hoka, and Puma's combined true lead over Nike. **(9 Points)**


```python
diff_lower_bound = ...
diff_upper_bound = ...
print("Bootstrapped 90% confidence interval for Adidas, Hoka, and Puma's combined true lead over Nike: [{:f}%, {:f}%]".format(diff_lower_bound, diff_upper_bound))
```

```python
grader.check("q1_6")
```

## 2. Interpreting Confidence Intervals 

The staff computed the following 90% confidence interval for the percentage of students who prefered Nike: 

$$[41.8, 46.27]$$

(Your answer may have been a bit different due to randomness; that doesn't mean it was wrong!)

<!-- BEGIN QUESTION -->

**Question 2.1.** The staff also created 70%, 95%, 99% confidence intervals from the same sample, but forgot to label which confidence interval represented which percentages. **First**, determine which confidence level (70%, 95%, or 99%) corresponds to the 99% confidence interval from the sample below. **Then**, explain your thought process and why you picked your answer **(10 Points)**

The intervals are below:

* [42.65, 45.62]
* [41.42, 46.81]
* [40.65, 47.96]




_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.2.** Suppose we produced 5,000 new samples (each one a new/distinct uniform random sample of 1,300 students) from the population and created a 90% confidence interval from each one. Roughly how many of those 5,000 intervals do you expect will actually contain the true percentage of the population? **(9 Points)**

Assign your answer to `true_percentage_intervals`.


```python
true_percentage_intervals = ...
```

```python
grader.check("q2_2")
```

Recall the second bootstrap confidence interval you created, which estimated Adidas, Hoka, and Puma combined lead over Nike. Among students in the sample, Nike was behind 8%. The staff's 90% confidence interval for the true lead (in the population of all voters) was:

$$[7.467283, 16.397229]$$

Suppose we are interested in testing a simple yes-or-no question:

> "Is the percentage of students who prefer Nike equal to the percentage of students who prefer Adidas, Hoka, and Puma combined?"

Our null hypothesis is that the percentages are equal, or equivalently, that Adidas, Hoka, and Puma's combined lead is exactly 0. Our alternative hypothesis is that Adidas, Hoka, and Puma's combined lead is not equal to 0.  In the questions below, don't compute any confidence interval yourself—use only the staff's 90% confidence interval.

**Question 2.3.** Say we use a 10% p-value cutoff. Do we reject the null, fail to reject the null, or are we unable to tell using the staff's confidence interval? **(9 Points)**

Assign `cutoff_ten_percent` to the number corresponding to the correct answer.

1. Fail to reject the null / Data is consistent with the null hypothesis
2. Reject the null / Data is consistent with the alternative hypothesis
3. Unable to tell using our staff confidence interval.

*Hint:* Consider the relationship between the p-value cutoff and confidence. If you're confused, take a look at [this chapter](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) of the textbook.


```python
cutoff_ten_percent = ...
```

```python
grader.check("q2_3")
```

**Question 2.4.** What if, instead, we use a p-value cutoff of 15%? Do we reject the null, fail to reject the null, or are we unable to tell using our staff confidence interval? **(9 Points)**

Assign `cutoff_fifteen_percent` to the number corresponding to the correct answer.

1. Fail to reject the null / Data is consistent with the null hypothesis
2. Reject the null / Data is consistent with the alternative hypothesis
3. Unable to tell using our staff confidence interval.


```python
cutoff_fifteen_percent = ...
```

```python
grader.check("q2_4")
```

**Question 2.5.** What if we use a p-value cutoff of 1%? Do we reject, fail to reject, or are we unable to tell using our confidence interval? **(9 Points)**

Assign `cutoff_one_percent` to the number corresponding to the correct answer.

1. Fail to reject the null / Data is consistent with the null hypothesis
2. Reject the null / Data is consistent with the alternative hypothesis
3. Unable to tell using our staff confidence interval.


```python
cutoff_one_percent = ...
```

```python
grader.check("q2_5")
```

You're done with Homework 8!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1_1 (points: [0, 0, 2])
```
>>> type(one_resampled_percentage(preferences)) in set([float, np.float64])
True
```
```
>>> one_resampled_percentage(preferences) >= 0
True
```
```
>>> 25 <= one_resampled_percentage(preferences) <= 50
True
```


Tests for q1_2 (points: [0])
```
>>> len(percentages_in_resamples()) == 2000
True
```


Tests for q1_3 (points: [0])
```
>>> 35 <= nike_lower_bound <= nike_upper_bound <= 50
True
```


Tests for q1_4 (points: [0])
```
>>> type(one_resampled_difference(preferences)) in set([float, np.float64])
True
```


Tests for q1_6 (points: [0])
```
>>> -5 <= diff_lower_bound <= diff_upper_bound <= 20
True
```


Tests for q2_2 (points: [0])
```
>>> 1000 <= true_percentage_intervals <= 6000
True
```


Tests for q2_3 (points: [0, 0, 0])
```
>>> 1 <= cutoff_ten_percent <= 3
True
```
```
>>> cutoff_ten_percent != 3
True
```
```
>>> cutoff_ten_percent != 1
True
```


Tests for q2_4 (points: [0])
```
>>> 1 <= cutoff_fifteen_percent <= 3
True
```


Tests for q2_5 (points: [0])
```
>>> 1 <= cutoff_one_percent <= 3
True
```
